<a href="https://colab.research.google.com/github/danial-amin/USCD-MicroMasters/blob/main/6_conditional_probability_hw.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

```python
import numpy as np
```

# Conditional Probability and Baye's Rule


## Problem

There are two urns $A$ and $B$. Urn $A$ contains $r_A$ red balls and $w_A$ white balls whereas urn $B$ contains $r_B$ red balls and $w_B$ white balls. One of the urns is picked at random and then one ball is picked at random from this urn. Write a function **conditional_probability** that calculates the conditional probability that the randomly chosen ball belonged to urn $A$ given that it is white. Assume that $\frac{r_A}{w_A}\neq\frac{r_B}{w_B}$.

<font  style="color:blue"> **Code**</font>
```python
rA, wA, rB, wB = 1., 2., 2., 1.
conditional__probability(rA, wA, rB, wB) 
```

<font  style="color:magenta"> **Output**</font>
```
0.6666666666666666
```

```python
# modify this cell

def conditional__probability(rA, wA, rB, wB):
   Prob_white_in_A= wA/(rA+wA)
   Prob_white = wA/(rA+wA) + wB/(rB+wB)
   return Prob_white_in_A/Prob_white

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)  
assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5) 

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

```python
conditional__probability( 2., 1., 6., 4.)







```

Output:
```
0.4545454545454545
```

```python

```

Output:
```
-870.2083333333333
0
-853.25
0
-836.4583333333334
0
-819.8333333333334
0
-803.375
0
-787.0833333333335
0
-770.9583333333334
0
-755.0
0
-739.2083333333335
0
-723.5833333333334
0
-708.125
0
-692.8333333333334
0
-677.7083333333333
0
-662.75
0
-647.9583333333334
0
-633.3333333333334
0
-618.875
0
-604.5833333333333
0
-590.4583333333334
0
-576.5
0
-562.7083333333333
0
-549.0833333333334
0
-535.625
0
-522.3333333333334
0
-509.20833333333337
0
-496.25
0
-483.4583333333333
0
-470.83333333333337
0
-458.375
0
-446.08333333333337
0
-433.95833333333337
0
-422.0
0
-410.20833333333337
0
-398.5833333333333
0
-387.125
0
-375.8333333333333
0
-364.7083333333333
0
-353.75
0
-342.9583333333333
0
-332.33333333333337
0
-321.875
0
-311.5833333333333
0
-301.45833333333337
0
-291.5
0
-281.7083333333333
0
-272.0833333333333
0
-262.625
0
-253.33333333333334
0
-244.20833333333334
0
-235.25000000000003
0
-226.45833333333331
0
-217.83333333333334
0
-209.375
0
-201.08333333333334
0
-192.95833333333337
0
-185.0
0
-177
```