<a href="https://colab.research.google.com/github/danial-amin/USCD-MicroMasters/blob/main/Prob%26StatinDS/Programming%20Assignments/3-sets_HW.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

*Note:* No packages should be imported for this assignment

# Sets: The Inclusion-Exclusion Principle

## Problem 1

### Problem 1.1

<img style="float: right;width:300px;height:200px" src="figs/Venn2.jpg">
The inclusion-exclusion principle states that for two sets $A$ and $B$,

$$|A\cup B|=|A|+|B|-|A\cap B|.$$

Write the following functions to determine $|A\cup B|$ in two different ways.

A function **union** that determines first $A\cup B$ and then evaluates the union's size.
Output the ordered pair $(A\cup B, |A\cup B|)$.

<font  style="color:blue"> * **Sample run** *</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
print union(A, B)
```

<font  style="color:magenta"> * **Expected Output** *</font>
```
({-6, 0, 1, 2, 3}, 5)
```

```python
# modify this cell

def union(A, B):
    # inputs: A and B are of type 'set'
    # output: a tuple of the type (set, set_length)
    return (A|B, len(A|B))

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
A = {-2,0,5, "bob"}
B = {-2,0,5,10,"jill"}
print(union(A,B))

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
({0, 'jill', 5, 10, 'bob', -2}, 6)
```

### Problem 1.2

A function **inclusion_exclusion** that first deterimines $|A|$, $|B|$, $A\cap B$, and $|A\cap B|$, and then uses the inclusion-exclusion formula to determine $|A\cup B|$. 
Output the tuple $(|A|, |B|, |A\cap B|, |A\cup B|)$.

<img style="float: right;width:400px;height:266px" src="figs/Venn3.jpg">
<font  style="color:blue"> * **Sample run:** *</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
print inclusion_exclusion(A, B)
print "notice: 3 + 4 - 2 == 5"
```

<font  style="color:magenta"> * **Expected Output:** *</font>
```
(3, 4, 2, 5)
notice: 3 + 4 - 2 == 5
```

```python
# modify this cell

def inclusion_exclusion(A, B):
    return (len(A),len(B),len(A&B),len(A)+len(B)-len(A&B))

```

```python
A = {-2,0,5, "bob"}
B = {-2,0,5,10,"jill"}
print(inclusion_exclusion(A,B))

```

Output:
```
(4, 5, 3, 6)
```

<!-- grade_id: ex2, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
assert inclusion_exclusion(A, B) == (5, 6, 2, 9)

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

## Problem 2

The inclusion-exclusion principle says that for three sets $A$, $B$ and $C$, 

$$|A\cup B\cup C|=|A|+|B|+|C|-|A\cap B|-|B\cap C|-|C\cap A|+|A\cap B\cap C|$$

We will write the following functions to determine $|A\cup B\cup C|$ in two different ways.

### Problem 2.1

Write function **union3** that first determines $A\cup B\cup C$ and then evaluates the size of this union.
Output the tuple $(A\cup B\cup C, |A\cup B\cup C|)$.

<font  style="color:blue"> * **Sample run:** *</font>
```python
A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
C = {2, 10}
union3(A, B, C)
```


<font  style="color:magenta"> * **Expected Output:** *</font>
```
({-6, 0, 1, 2, 3, 4, 5, 8, 9, 10}, 10)
```

```python
# modify this cell

def union3(A, B, C):
    # inputs: A, B and C are of type 'set'
    # output: a tuple of the type (set, set_length)
    return (A|B|C, len(A|B|C))
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: ex3, points: 5, grade: true, solution: false -->
```python
# check Function
A = {1, 2, 4, 5, 10}
B = {5, 2, -6, 5, 8, 9}
C = {2, 10, 13}
assert union3(A,B,C) == ({-6, 1, 2, 4, 5, 8, 9, 10, 13}, 9)

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

```python
A = {-2,0,5, "bob"}
B = {-2,0,5,10,"jill"}
C={-2,0,3,5,10}
print(union3(A,B,C))
```

Output:
```
({0, 'jill', 3, 5, 10, 'bob', -2}, 7)
```

### Problem 2.2

A function **inclusion_exclusion3** that first deterimines the sizes of $A$, $B$, $C$ and their mutual intersections, and then uses the inclusion-exclusion principle to determine the size of the union. Output the tuple $(|A\cap B\cap C|, |A\cup B\cup C|)$. Note that for brevity we are asking you to output the intermediate answer just for $A\cap B\cap C$, but you need to calculate all. 

<font  style="color:blue"> * **Sample run:** *</font>
```python
A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
C = {2, 10}
print inclusion_exclusion3(A, B, C)
```


<font  style="color:magenta"> * **Expected Output:** *</font>
```
(1, 10)
```

```python
# modify this cell

def inclusion_exclusion3(A, B, C):
    return (len(A&B&C),len(A)+len(B)+len(C)-len(A&B)-len(A&C)-len(B&C)+len(A&B&C))

```

```python
A = {-2,0,5, "bob"}
B = {-2,0,5,10,"jill"}
C={-2,0,3,5,10}
inclusion_exclusion3(A, B, C)
```

Output:
```
(3, 7)
```

<!-- grade_id: ex4, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {1, 2, 4, 5, 10}
B = {5, 2, -6, 5, 8, 9, 10}
C = {2, 10, 13}
assert inclusion_exclusion3(A,B,C) == (2, 9)

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```